---
  title: "Chapter 1"
  description: ""
---

## Variables

```yaml
type: NormalExercise 
xp: 100 
key: bda4ea2742   
```


We learn to define variables and print to screen in Python.

This is a variable assignment:
> `x = 5`


`@instructions`
Let's set the variables `x` and `y` to some number

`x = 5`

`y = 6`

let's sum them and store the result in `z`

`z = x + y`

and print `z` to screen

`print(z)`

`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}
x = 5
y = 6
z = x + y
print(z)
```

`@solution`

```{python}
x = 5
y = 6
z = x + y
print(z)
```

`@sct`

```{python}
1
```

---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 54ea0f3a94   
```


Test


`@instructions`
Test

`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}
a = 7
```

`@solution`

```{python}
a = 7
```

`@sct`

```{python}
1
```
