---
  title: "Test"
  description: "Test"
---

## Your first python script

```yaml
type: TabExercise 
xp: 100 
key: 83d1ec5467   
```


We will learn how to set up variables and print in Python.


`@pre_exercise_code`

```{python}
x = 5
y = 1
z = x + y
```

`@sample_code`

```{python}
print(z)
```

***



```yaml
type: NormalExercise 
xp: 100 
key: 8e68fb350c   
```





`@instructions`
Define two variables a and b and assign the values 6 and 7.
Assign the sum to c and print its value.

`@hint`


`@sample_code`

```{python}

```


`@solution`

```{python}
a = 6
b = 7
c = a + b
print(c)
```

`@sct`

```{python}

```


