---
  title: "Test"
  description: "Test"
---

## Variables

```yaml
type: NormalExercise 
xp: 100 
key: bda4ea2742   
```


We learn to define variables and print to screen in Python.


`@instructions`
`x = 5
y = 6
z = x + y
print(z)`

`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}
x =
y =
z =
print()
```

`@solution`

```{python}

```


`@sct`

```{python}

```

